---
title: Is is hotter or colder than normal?
description: >-
  Insert the chapter description here

---
## Mean or Median

```yaml
type: NormalExercise
key: 251207dc0e
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Correlations

```yaml
type: NormalExercise
key: 9fa53f15be
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
