---
title: Reading Data
description: >-
  Insert the chapter description here


---
## Importing the Data

```yaml
type: NormalExercise

xp: 

key: 3938a0ce13
```

We always begin the data science pipeline by asking a question. This question is what defines what we want to solve, or understand from the data.  However, we first need to take a see and explore our data and 

`@instructions`
We have three data files: 

- CR1000_OneHour.xlsx
- CR1000_OneMin.xlsx
- CR1000_FifteenSec.xlsx


Importing data in R is very simple when using prepackaged libraries.

`@hint`











---
## Viewing the last 10 Rows

```yaml
type: NormalExercise
skills: 1
lang: r
xp: 100
key: 98b8ed2d89
```














---
## General Information

```yaml
type: NormalExercise
key: 56a2ee92da
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
